# Lab: GPA 🎓

In this lab, you'll explore the GPA dataset and find out about the easiest, hardest, and most popular courses at Illinois!

The dataset you're using in this lab is the same dataset that Professor Wade uses to make the GPA visualizations! :)  Check this out here if you haven't seen it already: [https://waf.cs.illinois.edu/discovery/grade_disparity_between_sections_at_uiuc/](https://waf.cs.illinois.edu/discovery/grade_disparity_between_sections_at_uiuc/)

A few tips to remember:

- Refer to your **lecture notebook** and the **pandas cheat sheet** to help you out with the code!
- If you find yourself stuck for more than a few minutes, ask a neighbor or course staff for help!  When you're giving help to your neighbor, explain the **idea and approach** to the problem without sharing the answer itself so they can have the same **<i>ah-hah</i>** moment!
- We're here to help you!  Don't feel embarrassed or shy to ask us for help!

Let's get started!


<hr style="color: #DD3403;">

## Your Lab Group

Edit the next Python cell to add information about who you're working within your lab section:

```python
# First, meet your CAs and TA if you haven't already!
# First name is enough, we'll know who they are! :)
ta_name = ""
ca1_name = ""
ca2_name = ""


# Say hello to each other!
# - Groups of 3 are ideal :)
# - However, groups of 2 or 4 are fine too!
#
# QOTD to Ask Your Group: "Not including STAT 107, what is your favorite course this semester?"
partner1_name = ""
partner1_netid = ""
partner1_fav_course = ""

partner2_name = ""
partner2_netid = ""
partner2_fav_course = ""

partner3_name = ""
partner3_netid = ""
partner3_fav_course = ""
```

<hr style="color: #DD3403;">

## Puzzle 1: Basic Data Manipulation

### Puzzle 1.1: Loading the Dataset

For the first part of our lab, we'll be using the GPA dataset, which should be automatically downloaded when you fetch this notebook. 

In the following cell, load the pandas library and read the dataset into `df`.  You can find the GPA dataset at:

```
https://waf.cs.illinois.edu/discovery/gpa.csv
```

```python
# Import the gpa.csv dataset below

```

### Puzzle 1.2: Adding Helper Columns

Add a new column called `"Average GPA"`, that computes the average GPA for each course.  You've done this in the PL Homework, but now you're doing it for real! :)

- You shouldn't include students with a `"W"` in either column.
- The grade point scale can be found at https://registrar.illinois.edu/courses-grades/explanation-of-grades/.

<!-- grade_id: cell-1043285927ac8a85, points: 0, grade: false, solution: true -->
```python
# add the new column here
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 1.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, and RE-RUN your code and then this cell.
assert(abs(df[(df["YearTerm"] == "2018-fa") & (df["Subject"] == "STAT") & (df["Number"] == 100)]["Average GPA"].mean() - 3.247) < 0.01), "It appears that your \"Average GPA\" column is incorrect."

## == SUCCESS MESSAGE ==
# You'll only see this message (With the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<hr style="color: #DD3403;">

## Puzzle 2: Finding the Easiest and Hardest Courses

### Puzzle 2.1: Hardest Courses in the Dataset

Based on the `"Average GPA"` column, find the **30 hardest courses** at Illinois, i.e., the 30 courses with the lowest "Average GPA". Store the results in `df_hard`.

<!-- grade_id: cell-13f6d92b7a658526, points: 0, grade: false, solution: true -->
```python
df_hard = ...
df_hard
```

### Puzzle 2.2: "Level" of Hardest Courses

Find the average `"Number"` of the **30 hardest courses** at Illinois. Store the result as `df_hard_avg`.

<!-- grade_id: cell-a7106dca7694711c, points: 0, grade: false, solution: true -->
```python
df_hard_avg = ...
df_hard_avg
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-46729cc74c213e52, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.1 and 2.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell(s), make changes, and RE-RUN your code and then this cell.
import math
assert( math.isclose(df_hard_avg, 160.83333333333334) ), "It appears that your average \"Number\" is incorrect. Please check both Puzzle 2.1 and Puzzle 2.2."
## == SUCCESS MESSAGE ==
# You'll only see this message (With the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 2.3: ❓ Individual Reflection Question ❓ Are the Hardest Courses Really "Hard"?

After solving Puzzle 2.1 and 2.2, your friend has the following claim:

> *"We know that the undergraduate courses are coded from 001 to 499, where a larger number (in the hundreds place) usually implies more advanced material.  Based on our results in the previous puzzles, the data shows that the junior-level and senior-level courses are clearly not the hardest courses at UIUC."*

Comment on your friend's claim below. Do you think he is correct? Explain why or why not in at least three complete sentences.

*(✏️ Edit this cell to replace this text wth your answer this puzzle. ✏️)*

### Puzzle 2.4: Easiest Courses in 2021

Based on the `"Average GPA"` column, find the **30 easiest courses in 2021** at UIUC. Store the results in `df_easy_2021`.

<!-- grade_id: cell-44d17a2ef57b1a19, points: 0, grade: false, solution: true -->
```python
df_easy_2021 = ...
df_easy_2021
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-af77ae707b9fa776, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 2.4 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
import math
assert( len(df_easy_2021) == 30 ), "You should only have the 30 easiest courses in 2021."
assert( math.isclose(df_easy_2021["Average GPA"].mean(), 3.97873347181759) ), "Are you sure you have the 30 easiest courses in 2021?"
## == SUCCESS MESSAGE ==
# You'll only see this message (With the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

<hr style="color: #DD3403;">

## Puzzle 3: Finding Data by Subject

### Puzzle 3.1: Total Number of Grades by Subject

Find the total number of each grade ("A+"s, "A"s, "A-"s, etc.) given by different subjects. Store the results in `df_group`.

<!-- grade_id: cell-7d31cf501bfd8025, points: 0, grade: false, solution: true -->
```python
df_group = ...
df_group
```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-e9ad64dffa5a1cf1, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.1 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
assert( len(df_group) == 168 ), "Are you sure you grouped by subjects?"
assert( df_group[ df_group["Subject"] == "CS" ]["A"].values[0] == 60835 ), "Are you sure you counted the right number of each subject?"
## == SUCCESS MESSAGE ==
# You'll only see this message (With the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 3.2: Average GPA by Subject

In `df_group`, add a new column called `"Overall Average GPA"` that computes the overall GPA of each subject at UIUC.

<!-- grade_id: cell-7d3923eba40e1f45, points: 0, grade: false, solution: true -->
```python

```

### 🔬 Test Case Checkpoint 🔬

<!-- grade_id: cell-81b547cf0d42b67e, points: 1, grade: true, solution: false -->
```python
## == TEST CASES for Puzzle 3.2 ==
# - This read-only cell contains test cases for your previous cell.
# - If this cell runs without any error or output, you PASSED all test cases!
# - If this cell results in any errors, check your previous cell, make changes, RE-RUN your code, and then RE-RUN this cell.
import math
assert( math.isclose(df_group[df_group["Subject"] == "CS"]["Overall Average GPA"].values[0], 3.328369588478247) ), "Your overall average GPA doesn't appear to be correct."
## == SUCCESS MESSAGE ==
# You'll only see this message (With the emoji showing) if you passed all test cases:
tada = "\N{PARTY POPPER}"
print(f"{tada} All tests passed! {tada}")
```

### Puzzle 3.3: Visualizing the Grades of Subject Equal to `"STAT"`

Create a bar chart of the grades (`"A+"` to `"F"`) for the subject `"STAT"`.

```python

```

### Puzzle 3.4: ❓ Individual Reflection Question ❓ 

Bar Chart vs. Histogram

Why should we use a bar chart in Puzzle 3.3? Would it make sense to use a density histogram as well? Answer this question with at least 3 complete sentences!

*(✏️ Edit this cell to replace this text wth your answer this puzzle. ✏️)*

### Puzzle 3.5: Find Your Major

Finally, find your major in the dataset and check out how hard it is! 

Feel free to look around and explore! Part of being a data scientist is being creative and using Python to answer questions that you have.  Write your code below to determine how hard your major is and write comments in Puzzle 3.6 to explain what your code does. For example, what are the hardest classes in your major that have had at least 50 students?

<!-- grade_id: cell-116a2bb23ee81616, points: 0, grade: false, solution: true -->
```python
# Feel free to explore. Credit will be given as long as you try some new code (i.e., not directly copied from any puzzles).
```

### Puzzle 3.6: ❓ Individual Reflection Question ❓ 

Explain what you did in Puzzle 3.5 here, as well as what you found!  This is your chance to interpret your results and teach us something that you found interesting :) Write at least 3 complete sentences.

*(✏️ Edit this cell to replace this text wth your answer this puzzle. ✏️)*

<hr style="color: #DD3403;">

# Finale!

You're almost done! All you need to do is to commit your lab to GitHub:

1. Make certain to save your work. To do this, go to **File => Save All**

2. After you have saved, exit this notebook and follow the Canvas instructions to commit this lab to your Git repository!

3. Your TA will grade your submission and provide you feedback after the lab is due. :)